Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
IMMATRICULATION_NUMBER = ""
```

---

# Exercise 12: LZ78

Goal of this exercise is to implement the LZ78 compression.



## Compute the LZ78 encoding

Implement the method

<code>lz78 = encode(inputString)</code>
    
that gets the input String <code>inputString</code> and computes the LZ78 compressed output string simulated by a list of tuples. 
You do not need to write a binary String, but just fill a list of tuples consisting of id and next character c
You need the following steps to compress via LZ78:

* Search for the longest prefix of the String not yet encoded within the dictionary
* Append [i ,c] to the output list
* Add prefix (+) c to the dictionary

You do not need to implement the dictionary in form of a prefix trie, but can simply use a dictionary for that.

<!-- grade_id: cell-19abe5c1c5d59cb5, points: 0, grade: false, solution: true -->
```python
def encode(inputString):
    dictionary = {"": 0}
    output = []
    while inputString:
        prefix = ""
        for char in inputString:
            prefix += char
            if prefix not in dictionary:
                break
        if prefix in dictionary:
            # string completely consumed
            output.append([dictionary[prefix], ""])
            break
        dictionary[prefix] = len(dictionary)
        output.append([dictionary[prefix[:-1]], prefix[-1]])
        inputString = inputString[len(prefix):]
    return output


        
    
```

<!-- grade_id: cell-f0775ba64c1c34ea, points: 1, grade: true, solution: false -->
```python
assert encode('tobeornottobeortobeornot')== [[0, 't'],
 [0, 'o'],
 [0, 'b'],
 [0, 'e'],
 [2, 'r'],
 [0, 'n'],
 [2, 't'],
 [1, 'o'],
 [3, 'e'],
 [5, 't'],
 [2, 'b'],
 [4, 'o'],
 [0, 'r'],
 [6, 'o'],
 [1, '']]

assert encode('bed_spreaders_spread_spreads_on_beds') == [[0, 'b'],
 [0, 'e'],
 [0, 'd'],
 [0, '_'],
 [0, 's'],
 [0, 'p'],
 [0, 'r'],
 [2, 'a'],
 [3, 'e'],
 [7, 's'],
 [4, 's'],
 [6, 'r'],
 [8, 'd'],
 [11, 'p'],
 [7, 'e'],
 [0, 'a'],
 [3, 's'],
 [4, 'o'],
 [0, 'n'],
 [4, 'b'],
 [2, 'd'],
 [5, '']]
```

## Compute the LZ78 decoding

Implement the method

<code>decoded = decode(listOfTuples)</code>

The method should decode the given list of tuples and return the decoded String as result.


<!-- grade_id: cell-502fa970a1dfecd9, points: 0, grade: false, solution: true -->
```python
def decode(listOfTuples):
    dictionary = {0: ""}
    output = ""
    for index, char in listOfTuples:
        if index not in dictionary:
            raise Exception("Invalid input")
        output += dictionary[index] + char
        dictionary[len(dictionary)] = dictionary[index] + char
    return output
```

<!-- grade_id: cell-af3230b24d38d0b8, points: 1, grade: true, solution: false -->
```python
def test(inputString):
    assert decode(encode(inputString))==inputString, "Test failed for input string " + inputString
    
test('tobeornottobeortobeornot')
test('peterpiperpickedapeckofpickledpeppersapeckofpickledpepperspeterpiperpickedifpeterpiperpickedapeckofpickledpepperswheresthepeckofpickledpepperspeterpiperpicked')
test('iwishtowishthewishyouwishtowishbutifyouwishthewishthewitchwishesiwontwishthewishyouwishtowish')
test('howmanycookiescouldagoodcookcookifagoodcookcouldcookcookiesagoodcookcouldcookasmuchcookiesasagoodcookwhocouldcookcookies')

import random
import string

def get_random_string(length):
    letters = string.printable
    result_str = ''.join(random.choice(letters) for i in range(length))
    return result_str 
    
for i in range(10):
    randomS = get_random_string(10*(i+1))
    test(randomS)


```

```python

```